<a href="https://colab.research.google.com/github/PlaZMaD/GP_ML_2025/blob/main/Seminar_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Welcome

На практических занятиях курса мы будем использовать [язык программирования Python](https://www.python.org) в [среде Google Colab](https://colab.research.google.com). В качестве альтернативы вы можете загрузить какой-нибудь другой дистрибутив python, например [anaconda](https://www.anaconda.com/) и запустить jupyter локально (см. [docs](https://jupyter.readthedocs.io/en/latest/running.html)).

```python
!python --version
```

Если вы новичок в Python, пожалуйста, ознакомьтесь со следующим руководством:
 - https://docs.python.org/3.7/tutorial/

В частности, следующие части должны дать более или менее полное представление об основах, которые необходимо знать:
   - https://docs.python.org/3.7/tutorial/introduction.html
   - https://docs.python.org/3.7/tutorial/controlflow.html
   - https://docs.python.org/3.7/tutorial/datastructures.html
   - https://docs.python.org/3.7/tutorial/modules.html
   - https://docs.python.org/3.7/tutorial/classes.html

Не забывайте соблюдать [PEP-8](https://peps.python.org/pep-0008/). Вы также можете ознакомиться с другими [руководствами по стилю](https://google.github.io/styleguide/pyguide.html).

Обзор основных возможностей среды Google Colab можно найти в  [разделе](https://colab.research.google.com/notebooks/basic_features_overview.ipynb).

# Еще чуть-чуть Python'а
##Генераторы:
Генератор — это объект, который сразу при создании не вычисляет значения всех своих элементов.
Он хранит в памяти только последний вычисленный элемент, правило перехода к следующему и условие, при котором выполнение прерывается.
Доступ к объектам выполняется итеративно. Аналогичными интерфейсом доступа обладают практически все контейнеры.
Полезно уметь использовать это для модификации/создания контейнеров.

```python
a = range(5)
print(a)
for i in a:
  print (i)
```

```python
b = [i for i in a]
c = [i**2 for i in a if i%2==0]
d = {i:f"Квадратом числа {i} является {i**2}" for i in c}
print("b: ", b)
print("c: ", c)
print("d: ", d)
```

## Чуть-чуть о классах в Python
До этого мы работали только с готовыми типами объектов, но можно создавать и свои!

```python
class dog:
  legs = 4
  def __init__(self, name, age):
      self.name = name
      self.age = age
  def speak(self):
    print("Ауф!")
```

```python
Nemo = dog('Nemo', 6)
print(Nemo.name)
print(Nemo.age)
Nemo.speak()

Muhtar = dog('Muhtar', 10)
print(Muhtar.name)
print(Muhtar.legs)
```

```python
Muhtar.legs = 3
print(Muhtar.legs)
print(Nemo.legs)
```

```python
dog.legs = 2
print(Muhtar.legs)
print(Nemo.legs)
```

# Numpy and vectorized computing

Почти любая модель машинного обучения требует определенных вычислительных усилий, обычно связанных с решением задач линейной алгебры. К сожалению, сырой python ужасен в этом, поскольку каждая операция выполняется интерпретатором  во время запуска кода.

Поэтому вместо этого мы воспользуемся numpy - библиотекой, позволяющей выполнять молниеносные вычисления с векторами, матрицами и другими тензорами. Основным объектом здесь является numpy.ndarray:

**please keep running all the code cells as you read**

```python
import numpy as np

a = np.array([1,2,3,4,5])
b = np.array([5,4,3,2,1])
print("a = ", a)
print("b = ", b)

# math and boolean operations can applied to each element of an array
print("a + 1 =", a + 1)
print("a * 2 =", a * 2)
print("a == 2", a == 2)
# ... or corresponding elements of two (or more) arrays
print("a + b =", a + b)
print("a * b =", a * b)
```

<!-- grade_id: cell-c8488cc1b2d24a07, points: 0, grade: false, solution: true -->
```python
# Your turn: compute half-products of a and b elements (halves of products)

def half_product(a, b):
    # your code here
    raise NotImplementedError
```

Убедитесь, что автоматические проверки пройдены!


<!-- grade_id: HalfProduct, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(
    half_product(
        np.array([ 1.3549439 , -1.69765972,  0.32111273, -0.2703243 , -0.38928596]),
        np.array([-1.50693732,  0.69443195,  0.90948151,  1.1717395 ,  1.4069652 ])
    ),
    np.array([-1.02090776, -0.58945457,  0.14602305, -0.15837483, -0.2738559 ])
)

print("a =", a)
print("b =", b)
print("a * b / 2 = ", half_product(a, b))
```

<!-- grade_id: cell-68ccf91f62d2ac06, points: 0, grade: false, solution: true -->
```python
# compute elementwise quotient between squared a and (b plus 1), return the result

def a_squared_over_b_plus_1(a, b):
    # your code here
    raise NotImplementedError
```

Убедитесь, что автоматические проверки пройдены!

<!-- grade_id: ASquaredOverBPlus1, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(
    a_squared_over_b_plus_1(
        np.array([ 1.3549439 , -1.69765972,  0.32111273, -0.2703243 , -0.38928596]),
        np.array([-1.50693732,  0.69443195,  0.90948151,  1.1717395 ,  1.4069652 ])
    ),
    np.array([-3.62149895,  1.70089364,  0.05400072,  0.03364825,  0.06296043])
)

print("a = ", a)
print("b = ", b)
print("a**2 / (b + 1)", a_squared_over_b_plus_1(a, b))
```

---------

Существует множество функций для создания массивов нулей, единиц, возрастающих чисел и т.д.:

```python
np.zeros(shape=(3, 4))
```

```python
np.ones(shape=(2, 5))
```

```python
np.arange(3, 15, 2.5) # start, stop, step
```

```python
np.linspace(0, 10, 11) # divide [0, 10] interval into 11 points
```

Вы можете легко изменять форму массивов:

```python
np.arange(24).reshape(2, 3, 4)
```

Размеры массива автоматически транслируются при выполнении операций:

```python
np.arange(3).reshape(1, 3) - np.arange(3).reshape(3, 1)
```

Вы можете использовать трансляцию различными способами. Например, для поэлементного возведения матрицы в степень:

```python
A = np.array([
    [1, 2],
    [3, 5]
])
powers = np.arange(1, 5)

A.reshape(1, 2, 2)**powers.reshape(4, 1, 1)
```

Существует также несколько способов объединения массивов. Например, `np.concatenate` объединяет массивы вдоль **существующей** оси:

```python
matrix1 = np.arange(18).reshape(6, 3) # array of shape (6, 3)
matrix2 = -np.arange(12).reshape(6, 2) # array of shape (6, 2)

np.concatenate([matrix1, matrix2], axis=1) # result of shape (6, 5)
```

В то время как np.stack добавляет новое измерение к результату:

```python
np.stack([
    np.arange(5), # array of shape (5,)
    np.linspace(0, 1, 5) # array of shape (5,)
], axis=1) # result of shape (5, 2)
```

Любая матрица может быть легко транспонирована:

```python
print(matrix2)
print('---')
print(matrix2.T)
```

```python
print('matrix2.shape =', matrix2.shape)
print('matrix2.T.shape =', matrix2.T.shape)
```

<!-- grade_id: cell-7052c430a2292e3f, points: 0, grade: false, solution: true -->
```python
# Your turn!
# Using the methods you've learned (array initialization,
# reshaping and concatenation), write a function that composes
# and returns the following matrix:
#
#              0   7  14  21  28  0  0
#              1   8  15  22  29  0  0
#              2   9  16  23  30  0  0
#              3  10  17  24  31  3  3
#              4  11  18  25  32  3  3
#              5  12  19  26  33  3  3
#              6  13  20  27  34  3  3


def compose_matrix():
    # your code here
    raise NotImplementedError
```

<!-- grade_id: ComposeMatrix, points: 1, grade: true, solution: false -->
```python
np.testing.assert_equal(
    compose_matrix(),
    np.array([[0,   7,  14,  21,  28,  0,  0],
              [1,   8,  15,  22,  29,  0,  0],
              [2,   9,  16,  23,  30,  0,  0],
              [3,  10,  17,  24,  31,  3,  3],
              [4,  11,  18,  25,  32,  3,  3],
              [5,  12,  19,  26,  33,  3,  3],
              [6,  13,  20,  27,  34,  3,  3]])
)
```

Линейная алгебра:

```python
a = np.arange(3)
B = np.arange(12).reshape(4, 3)
print("a =", a)
print("B =\n", B)
print("Dot product (a*a):", a @ a) # or: np.dot(a, a)
print("Matrix-vector (B*a):", B @ a) # or: np.matmul(B, a)
print("Matrix-matrix (B*B.T):\n", B @ B.T) # or: np.matmul(B, B.T)
```

И огромное колиество прочих математических операций:

```python
a = np.array([1,2,3,4,5])
b = np.array([5,4,3,2,1])
print("numpy.sum(a) = ", np.sum(a))
print("numpy.mean(a) = ", np.mean(a))
print("numpy.min(a) = ",  np.min(a))
print("numpy.argmin(b) = ", np.argmin(b))  # index of minimal element
print("numpy.dot(a,b) = ", np.dot(a, b))      # dot product. Also used for matrix/tensor multiplication
print("numpy.unique(['male','male','female','female','male']) = ", np.unique(['male','male','female','female','male']))

# and tons of other stuff. see http://bit.ly/2u5q430 .
```

```python
# most of this functions are also implemented as members of numpy arrays, e.g.:
print('a.min() =', a.min())
print('a.mean() =', a.mean())
```

# Pandas -  работа с табличными данными

Pandas - это библиотека, которая поможет вам загрузить данные, подготовить их и провести легкий анализ. В качестве основного объекта здесь выступает pandas.DataFrame - двумерная таблица  (на самом деле под капотом работает numpy).


Загрузим наш датасет:

```python
!wget https://raw.githubusercontent.com/HSE-LAMBDA/MLDM-2022/main/01-intro/tabular-playground-series-aug-2022.zip
```

```python
!unzip /content/tabular-playground-series-aug-2022.zip -d data
```

Теперь давайте используем Pandas для чтения данных

```python
import pandas as pd
data = pd.read_csv('/content/data/train.csv',
                    index_col='id') # this yields a pandas.DataFrame
data.index += 1
```

```python
# Take a look at the data

data.head(5) # selects top 5 lines
```

```python
pd.options.display.max_rows = 10
data
```

```python
data.sample(5)
```

#### О данных
Эти данные представляют собой результаты большого исследования по тестированию продукции. Для каждого `кода_продукта` вам дается ряд `атрибутов` продукта (фиксированных для кода), а также ряд `измерений` для каждого отдельного продукта, представляющих различные методы лабораторных испытаний. Каждый продукт используется в эксперименте, смоделированном в реальной среде, и поглощает определенное количество жидкости (`нагрузка`), чтобы увидеть, не выйдет ли он из строя.

Ваша задача - использовать эти данные для прогнозирования `failure` для новых кодов с учетом результатов их лабораторных испытаний. Обучающие данные содержат целевое значение `failure`, и вам нужно предсказать вероятность того, что каждый испытуемый образец получит `failure`.

```python
# table dimensions
print("len(data) = ", len(data.index))
print("data.shape = ", data.shape)
```

```python
# select a single row
print(data.loc[4])
```

```python
# select a single column.
loadings = data["loading"] # alternatively: data.loading
print(loadings.loc[:10])
```

```python
# select several columns and rows at once
data.loc[5:10, ("loading", "failure")]    # alternatively: data[["loading","failure"]].loc[5:10]
```

### `loc` vs `iloc`

В pandas есть два способа индексирования строк:

*  по значениям столбцов и индекса (в нашем случае id) - для этого используется data.loc
*  по позиционному индексу - для этого используется data.iloc

Обратите внимание, что индексный столбец начинается с 1, поэтому позиционный индекс 0 будет соответствовать значению индексного столбца 1, позиционный 1 - значению индексного столбца 2, и так далее:

```python
print(data.index)
print('------')
print("data.iloc[0]:")
print(data.iloc[0])
print('------')
print("data.loc[1]:")
print(data.loc[1])
```

Также обратите внимание, что при индексировании с помощью `.loc` включаются оба конца среза:

```python
data.loc[2:3]
```

в то время как при использовании `.iloc` конец исключается:

```python
data.iloc[1:2]
```

Более сложная индексация (похожая на булевую индексацию в numpy):

```python
data.loc[(data['loading'] < 90) & (data['product_code'] == np.random.choice(data.product_code.unique()))]
```

```python
data.query('2.1 * measurement_9 < measurement_10 and loading < 111')
```

### Your turn:


```python
# select studies number 13 and 666 in a single line - did they fail?

print(data.loc[[13, 666], 'failure'])
```

```python
# compute the overall fail-rate (what fraction of studies failed)
# do we face a balanced train dataset?

print(1 - (data['failure'].sum())/len(data.index))
```

В Pandas также есть некоторые базовые инструменты анализа данных. Например, вы можете быстро вывести статистические показатели для каждого столбца с помощью `.describe()`.

```python
data.describe()
```

Некоторые столбцы содержат значения __NaN__ - это означает, что в них нет данных. Например, в исследовании `#26565` неизвестно значение `measurement_16`. Чтобы упростить дальнейший анализ данных, мы заменим значения NaN с помощью функции pandas `fillna`.

Примечание: мы делаем это, потому что это учебный материал. В общем случае, подумайте дважды, прежде чем изменять данные подобным образом.

```python
data['measurement_9'] = data['measurement_9'].fillna(value=data['measurement_9'].mean())
data['measurement_16'].fillna(value=data['measurement_16'].mean(), inplace = True)
```

```python
data.iloc[26565]
```

### Pandas + numpy

Важная часть: поскольку pandas использует numpy под капотом, большинство функций numpy работает с фреймами данных, поскольку вы можете получить их представление в numpy с помощью `.values` (большинство функций numpy будут работать даже на чистых объектах pandas):

```python
# calling np.max on a pure pandas column:
column_name = 'measurement_17'
print("Max {}: ".format(column_name), np.max(data[column_name]))

# calling np.argmax on a numpy representation of a pandas column
# to get its positional index:
print("\nThe study with the max " + column_name + ":\n",
      data.iloc[
          np.argmax(data[column_name].values)
      ])
```

```python
# numpy works only with positional index:
column_name = 'measurement_16'
print(data[column_name].values.argmax())
#     ^^^^^^^^^^^^^^^^^^^
#     this part returns a numpy array, argmax of which we are calculating


# in pandas you can ask for the index (i.e. value of the index column)
# of the maximal element like this:
print(data[column_name].idxmax())
```

### Your turn

Use numpy and pandas to answer a few questions about data

```python
# your code: compute mean loading and find max one.

#YOUR CODE
```

```python
# which product code is more likely to fail?
prod_names = data.product_code.unique()
print(prod_names)
#YOUR CODE
```

# Графики

В настоящее время наиболее актуальным выбором являются библиотеки Matplotlib и Plotly.

```python
import matplotlib.pyplot as plt
%matplotlib inline

import plotly.graph_objects as go
```

```python
n_points = 5
x = np.arange(n_points)
y = x**2

print("x =", x)
print("x**2 =", y)
print("plotting x**2 vs x:")


plt.scatter(x, y)
plt.show()  # show the first plot to begin drawing the next one

fig = go.Figure()
fig.add_trace(go.Scatter(x=x, y=y, mode='markers'))
fig.show()
```

```python
# histogram - showing data density
float_cols = [c for c in data.columns if data[c].dtype == float]

_, axs = plt.subplots(2, 4, figsize=(10,5))

for f, ax in zip(float_cols[:8], axs.ravel()):
    bins = np.linspace(min(data[f]), max(data[f]), 50)
    ax.hist(data[f], bins=bins, density=True)
    ax.set_xlabel(f)

plt.tight_layout(w_pad=1)
plt.suptitle('Distributions of the continuous features', fontsize=15, y=1.12)
plt.show()
```

```python
# .plot() method allows you to plot in different styles
data.product_code.value_counts().plot(kind='pie', autopct="%1.1f%%",shadow=True,
        startangle=45, explode=[0.065] * data.product_code.nunique());
```

#Машинное обучение c  scikit-learn

Scikit-learn - это инструмент для создания простых конвейеров машинного обучения.

Это единая библиотека, которая объединяет множество моделей под общим интерфейсом:

Create:__ model = sklearn.whatever.ModelNameHere(parameters_if_any)__.

Train:__ model.fit(X,y)__

Predict:__ model.predict(X_test)__.

В нем также есть утилиты для извлечения признаков, оценки качества или кросс-валидации.


#K-nearest neighbors

```python
from sklearn.neighbors import KNeighborsClassifier
from sklearn.metrics import accuracy_score
from sklearn.model_selection import train_test_split


x = data[float_cols].fillna(0).copy()
y = data["failure"]

model = KNeighborsClassifier(n_neighbors=5)

# split the data into train(90%) and test(10%)
X_train, X_test, y_train, y_test = train_test_split( x, y, test_size=0.1, random_state=42)


# fit the model
model.fit(X_train,  y_train)

# make the prediction
test_predictions = model.predict(X_test)
print("Test accuracy:", accuracy_score(y_test, test_predictions))

```

# Линейная регрессия

## y = kx + b

Начнем с одномерного случая, где истиная зависимость будет иметь вид:
$$y=k\cdot x+b+\text{noise}$$

```python
def linear_function(x):
  return 0.33 * x + 8.3

def gen_dataset(N, func, lims=(-1., 1.), noise_lvl=0.2):
  x = np.random.uniform(*lims, size=N)
  y = func(x) + noise_lvl * np.random.normal(size=x.shape)
  return x[:,None], y

X, y = gen_dataset(50, linear_function)
x = np.linspace(-1, 1, 101)
plt.plot(x, linear_function(x))
plt.scatter(X, y);
```

```python
from sklearn.linear_model import LinearRegression

model = LinearRegression()

model.fit(X, y)

x = np.linspace(-1, 1, 101)
plt.plot(x, linear_function(x), label='true function')
plt.scatter(X, y);
plt.plot(x, model.predict(x[:,None]), label='prediction')
plt.legend()

print(model.coef_, model.intercept_)
```

А что если случай более сложный?

```python
def true_function(x):
  return np.sin(3 * x + 0.8) + np.sin(1. / (x + 1.23))
```

```python
model = LinearRegression()
X, y = gen_dataset(25, true_function)

model.fit(X, y)

x = np.linspace(-1, 1, 101)
plt.plot(x, true_function(x), label='true function')
plt.scatter(X, y);
plt.plot(x, model.predict(x[:,None]), label='prediction')
plt.legend();
```

На самом деле, даже довольно сложные задачи можно решать простыми алгоритмами, если есть глубокое понимание этой самой задачи!
В данном случае может помочь конструирование искуственных фичей!
Для удобства работы мы будем цепочку обработки данных помещать в pipeline.


$$X'=
\begin{pmatrix}
x_1 & (x_1)^2 & \ldots & (x_1)^k \\
x_2 & (x_2)^2 & \ldots & (x_2)^k \\
\vdots & \vdots & \ddots & \vdots \\
x_N & (x_N)^2 & \ldots & (x_N)^k
\end{pmatrix},
$$

```python
from sklearn.preprocessing import PolynomialFeatures
from sklearn.pipeline import make_pipeline

poly_expand = PolynomialFeatures(3)
poly_expand.fit_transform(np.arange(5)[:,None])
```

```python
# The first parameter is the power of expansion. Try playing around with it.
poly_expand = PolynomialFeatures(1, include_bias=False)
linear_model = LinearRegression()
model = make_pipeline(
    poly_expand, linear_model
)

model.fit(X, y)

x = np.linspace(-1, 1, 101)
plt.plot(x, true_function(x), label='true function')
plt.scatter(X, y);
plt.plot(x, model.predict(x[:,None]), label='prediction')
plt.ylim(y.min() - 0.5, y.max() + 0.5)
plt.legend();
```

# Рассмотрим более сложный датасет - Boston Housing

 Variables in order:

 CRIM     per capita crime rate by town

 ZN       proportion of residential land zoned for lots over 25,000 sq.ft.

 INDUS    proportion of non-retail business acres per town

 CHAS     Charles River dummy variable (= 1 if tract bounds river; 0 otherwise)

 NOX      nitric oxides concentration (parts per 10 million)

 RM       average number of rooms per dwelling

 AGE      proportion of owner-occupied units built prior to 1940

 DIS      weighted distances to five Boston employment centres

 RAD      index of accessibility to radial highways

 TAX      full-value property-tax rate per $10,000

 PTRATIO  pupil-teacher ratio by town

 B        1000(Bk - 0.63)^2 where Bk is the proportion of blacks by town

 LSTAT    % lower status of the population

 MEDV     Median value of owner-occupied homes in $1000's

```python
!wget https://raw.githubusercontent.com/selva86/datasets/master/BostonHousing.csv
```

```python
raw_df = pd.read_csv('BostonHousing.csv')
print(raw_df.head())

y = raw_df['medv']
X = raw_df.drop('medv', axis=1)
```

Мы хотим предсказать MEDV - цену жилья, исходя из остальных показателей. Оценим зависимость целевой переменной от фичей:

```python
plt.figure(figsize=(15, 15))

grid_size = int(np.ceil(X.shape[1]**0.5))

for i, (name, x) in enumerate(X.items(), 1):
  plt.subplot(grid_size, grid_size, i)
  plt.scatter(x, y)
  plt.xlabel(name)
```

```python
from sklearn.linear_model import LinearRegression
from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import PolynomialFeatures
from sklearn.metrics import mean_squared_error
```

Выберем 3 наиболее удачных фичи и порстроим модель с их использованием.

```python
old_columns = X.columns
print(old_columns)
X.columns = [column.upper() for column in X.columns]
```

```python
columns = ["CRIM", "RM", "LSTAT"]

X_subset = X[columns]
X_subset /= X_subset.max()
```

```python
model = make_pipeline(
    PolynomialFeatures(9, include_bias=False), # Can you calculate how many features this will result in? :)
    LinearRegression()
)

model.fit(X_subset, y)
print('mse = ', mean_squared_error(y, model.predict(X_subset)))

plt.figure(figsize=(15, 4))
for i, c in enumerate(columns, 1):
  plt.subplot(1, len(columns), i)
  plt.scatter(X[c], y, label='data')
  plt.scatter(X[c], model.predict(X_subset), label='prediction')
  plt.legend()
```

На сколько это хорошо? На вид здорово, а что на самом деле?

# Train-Test split

```python
from sklearn.model_selection import train_test_split
```

```python
X_train, X_test, y_train, y_test = train_test_split(X_subset, y, test_size=50, random_state=39)

model = make_pipeline(
    PolynomialFeatures(9, include_bias=False),
    LinearRegression()
)

model.fit(X_train, y_train)

print('train mse = ', mean_squared_error(y_train, model.predict(X_train)))
print('test mse = ', mean_squared_error(y_test, model.predict(X_test)))
```

WAT?!

```python
plt.figure(figsize=(15, 4))
for i, c in enumerate(columns, 1):
  plt.subplot(1, len(columns), i)
  plt.scatter(X_test[c], y_test, label='data')
  plt.scatter(X_test[c], model.predict(X_test), label='prediction')
  plt.legend()
```

Может что-то не то с коэффициентами?

```python
model.named_steps['linearregression'].coef_
```

# L2 regularization (ridge regression)

Let's regularize the solution!

```python
from sklearn.linear_model import Ridge
```

```python
model = make_pipeline(
    PolynomialFeatures(9, include_bias=False),
    Ridge(alpha=1.)
)

model.fit(X_train, y_train)

print('train mse = ', mean_squared_error(y_train, model.predict(X_train)))
print('test mse = ', mean_squared_error(y_test, model.predict(X_test)))
```

```python
plt.figure(figsize=(15, 4))
for i, c in enumerate(columns, 1):
  plt.subplot(1, len(columns), i)
  plt.scatter(X_test[c], y_test, label='data')
  plt.scatter(X_test[c], model.predict(X_test), label='prediction')
  plt.legend()
```

```python
model.named_steps['ridge'].coef_
```

Now we'll study how losses and parameter values depend on the regularization power.

```python
from tqdm import tqdm
```

```python
reg_powers = np.logspace(-12, 5, 18 * 5, base=10)


train_mse = []
test_mse = []

params = []

for alpha in tqdm(reg_powers):
  linear_model = Ridge(alpha=alpha)
  model = make_pipeline(
    PolynomialFeatures(9, include_bias=False),
    linear_model
  )
  model.fit(X_train, y_train)

  params.append(
      np.append(linear_model.coef_,
                linear_model.intercept_)
  )

  train_mse.append(mean_squared_error(y_train, model.predict(X_train)))
  test_mse.append(mean_squared_error(y_test, model.predict(X_test)))

params = np.array(params)


plt.figure(figsize=(12, 5))

plt.subplot(1, 2, 1)

plt.plot(1. / reg_powers, train_mse, label='train')
plt.plot(1. / reg_powers, test_mse, label='test')
plt.ylabel('MSE')
plt.xlabel('inverse regularization power')
plt.legend()
plt.xscale('log')
plt.yscale('log')

plt.subplot(1, 2, 2)
plt.plot(1. / reg_powers, np.abs(params));
plt.xlabel("inverse regularization power")
plt.ylabel("parameter magnitude")
plt.xscale('log')
plt.yscale('log')
```

# L1 regularization (lasso regression)

Here's a similar study with the Lasso regression:

```python
from sklearn.linear_model import Lasso
```

```python
reg_powers = np.logspace(-4, 1, 6 * 5, base=10)

train_mse = []
test_mse = []

params = []

for alpha in tqdm(reg_powers):
  # Lasso doesn't have an analytic solution. Instead it
  # utilizes an iterative procedure, which for small
  # alpha values may take a while to converge.
  linear_model = Lasso(alpha=alpha, max_iter=1000000)
  model = make_pipeline(
    PolynomialFeatures(9, include_bias=False),
    linear_model
  )
  model.fit(X_train, y_train)

  params.append(
      np.append(linear_model.coef_,
                linear_model.intercept_)
  )

  train_mse.append(mean_squared_error(y_train, model.predict(X_train)))
  test_mse.append(mean_squared_error(y_test, model.predict(X_test)))

params = np.array(params)

plt.figure(figsize=(18, 5))

plt.subplot(1, 3, 1)

plt.plot(1. / reg_powers, train_mse, label='train')
plt.plot(1. / reg_powers, test_mse, label='test')
plt.ylabel('MSE')
plt.xlabel('inverse regularization power')
plt.legend()
plt.xscale('log')

plt.subplot(1, 3, 2)
plt.plot(1. / reg_powers, np.abs(params));
plt.xlabel("inverse regularization power")
plt.ylabel("parameter magnitude")
plt.xscale('log')
plt.yscale('log')

plt.subplot(1, 3, 3)
plt.plot(1. / reg_powers, np.isclose(params, 0.).sum(axis=1));
plt.xlabel("inverse regularization power")
plt.ylabel("# zero parameters")
plt.xscale('log')
```

#Самостоятельная работа:
Для отработки пройденного материала очень полезно проверить свои силы на отдельном задании:
https://github.com/PlaZMaD/GP_ML_2025/blob/main/Seminar_2_HW.ipynb

В этот раз проверка производиться не будет, все остается на вашей совести.

